## Codio Activity 7.1: Fitting a Simple Regression Line Usign Plotly and Sklearn

**Expected Time: 60 Minutes**

**Total Points: 20 Points**

This activity focuses on using `sklearn` to build a `LinearRegression` estimator.  For the dataset, another built-in seaborn dataset with information on geyser explosions is used.  Using this dataset, you are to build a regression model using the wait time to predict the duration of the explosion.

## Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)
- [Problem 4](#Problem-4)

```python
import matplotlib.pyplot as plt
import seaborn as sns
import numpy as np

import pandas as pd
from sklearn.linear_model import LinearRegression
```

### The Geyser Data

The dataset contains information on the waiting time for a geyser explosion, the duration of the explosion, and a categorization of the explosion duration.  This data comes from the seaborn built-in datasets.

```python
geyser = sns.load_dataset('geyser')
```

```python
geyser.head()
```

Output:
```
duration  waiting   kind
0     3.600       79   long
1     1.800       54  short
2     3.333       74   long
3     2.283       62  short
4     4.533       85   long
```

[Back to top](#Index:) 

## Problem 1

### Declaring `X` and `y`.  

**5 Points**

Assign the column `waiting` as a DataFrame to the variable `X` and the column `duration` as a Series to the variable `y` below.  

<!-- grade_id: cell-b8bc4f8de3b17e42, points: 0, grade: false, solution: true -->
```python
### GRADED

X = ''
y = ''

### BEGIN SOLUTION
X = geyser[['waiting']]
y = geyser['duration']
### END SOLUTION

# Answer check
print(X.shape, y.shape)
print(type(X), type(y))
```

Output:
```
(272, 1) (272,)
<class 'pandas.core.frame.DataFrame'> <class 'pandas.core.series.Series'>
```

<!-- grade_id: cell-8483d4251dee4b5c, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
X_ = geyser[['waiting']]
y_ = geyser['duration']
#
#
#
assert type(X_) == type(X)
pd.testing.assert_series_equal(y, y_)
pd.testing.assert_frame_equal(X_, X)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 2

### Building a model with `LinearRegression` 

**5 Points**

Below, instatiate and fit a regression model using your `X` and `y` from above.  Make sure to use only the default settings.  Assign your regressor to the variable `linreg` below.  

<!-- grade_id: cell-94aea8297583c01a, points: 0, grade: false, solution: true -->
```python
### GRADED

linreg = ''

### BEGIN SOLUTION
linreg = LinearRegression().fit(X, y)

### END SOLUTION

# Answer check
print(linreg)
```

Output:
```
LinearRegression()
```

<!-- grade_id: cell-05a1d7a647b68095, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
linreg_ = LinearRegression().fit(X, y)
coefs_ = linreg_.coef_
intercept_ = linreg.intercept_
coefs =  linreg.coef_
intercept = linreg.intercept_
#
#
#
assert type(linreg_) == type(linreg)
np.testing.assert_array_equal(coefs_, coefs)
np.testing.assert_array_equal(intercept, intercept_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 3

### Adding a prediction column

**5 Points**

Add a column of predictions to the geyser DataFrame named `prediction`.  Check that your DataFrame geyser contains the new column. 


<!-- grade_id: cell-2cc296d4cf53bb86, points: 0, grade: false, solution: true -->
```python
### GRADED

geyser['prediction'] = ''

### BEGIN SOLUTION
geyser['prediction'] = linreg.predict(X)

### END SOLUTION

# Answer check
print(geyser.columns)
print(geyser.shape)
```

Output:
```
Index(['duration', 'waiting', 'kind', 'prediction'], dtype='object')
(272, 4)
```

<!-- grade_id: cell-06970840d6ded5ba, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
geyser_ = sns.load_dataset('geyser')
X_ = geyser_[['waiting']]
y_ = geyser_['duration']
linreg_ = LinearRegression().fit(X_, y_)
geyser_['prediction'] = linreg_.predict(X_)
#
#
#
assert geyser.shape == geyser_.shape
pd.testing.assert_frame_equal(geyser, geyser_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 4

### Equation of line

**5 Points**

The equation of the line will be of the form

$$\text{duration} = \text{waiting}\times \text{slope} + \text{intercept}$$

Assign the slope of the solution as a float correct to two decimal places to `slope` below.  Also, assign the intercept as a float correct to two decimal places to `intercept` below.  

<!-- grade_id: cell-86e8eb5cdc55a423, points: 0, grade: false, solution: true -->
```python
### GRADED

slope = ''
intercept = ''

### BEGIN SOLUTION
slope = round(float(linreg.coef_), 2)
intercept = round(float(linreg.intercept_), 2)

### END SOLUTION

# Answer check
print(type(slope))
print(slope, intercept)
```

Output:
```
<class 'float'>
0.08 -1.87
```

<!-- grade_id: cell-c78c6f171e5a619e, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
geyser_ = sns.load_dataset('geyser')
X_ = geyser_[['waiting']]
y_ = geyser_['duration']
linreg_ = LinearRegression().fit(X_, y_)
intercept_ = round(float(linreg_.intercept_), 2)
slope_ = round(float(linreg_.coef_), 2)
#
#
#
assert type(slope) == type(slope_)
assert slope == slope_, 'Make sure the slope is rounded'
assert intercept == intercept_, 'Make sure the intercept is rounded'
### END HIDDEN TESTS
```