### Codio Activity 8.5: Comparing Complexity and Variance

**Expected Time: 60 Minutes**

**Total Points: 35**

In this activity, you will explore the effect of model complexity on the variance in predictions.  Continuing with the automotive data, you will build models on a subset of 10 vehicles.  You will compare the model error when used on the entire dataset, and investigate how variance changes with model complexity.

#### Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)
- [Problem 4](#Problem-4)


```python
import numpy as np
import pandas as pd
from sklearn.linear_model import LinearRegression
from sklearn.preprocessing import PolynomialFeatures
from sklearn.pipeline import Pipeline
from sklearn.metrics import mean_squared_error
import plotly.express as px
```

```python
auto = pd.read_csv('data/auto.csv')
```

```python
auto.head()
```

Output:
```
mpg  cylinders  displacement  horsepower  weight  acceleration  year  \
0  18.0          8         307.0       130.0    3504          12.0    70   
1  15.0          8         350.0       165.0    3693          11.5    70   
2  18.0          8         318.0       150.0    3436          11.0    70   
3  16.0          8         304.0       150.0    3433          12.0    70   
4  17.0          8         302.0       140.0    3449          10.5    70   

   origin                       name  
0       1  chevrolet chevelle malibu  
1       1          buick skylark 320  
2       1         plymouth satellite  
3       1              amc rebel sst  
4       1                ford torino
```

### The Sample

Below, a sample of ten vehicles from the data is extracted.  These data are what will form our **training** data.  The data is subsequently split into `X_train` and `y_train`.  You are to use this smaller dataset to build your models on, and explore their performance using the entire dataset.

```python
X = auto.loc[:,['horsepower']]
y = auto['mpg']
sample = auto.sample(10, random_state = 22)
X_train = sample.loc[:, ['horsepower']]
y_train = sample['mpg']
```

```python
X_train
```

Output:
```
horsepower
280        88.0
57         80.0
46        100.0
223       110.0
303        90.0
73        140.0
98        100.0
250       105.0
254       100.0
337       110.0
```

```python
y_train
```

Output:
```
280    22.3
57     25.0
46     19.0
223    17.5
303    28.4
73     13.0
98     18.0
250    19.2
254    20.5
337    23.5
Name: mpg, dtype: float64
```

```python
X.shape
```

Output:
```
(392, 1)
```

[Back to top](#Index:) 

### Problem 1

#### Iterate on Models

**20 Points**

In this problem, you are to again build models using degree 1 through 10.  Use a `Pipeline` and be sure to set `include_bias = False` in your transformer.  Fit your pipelines on the training data, and assign the predictions using the entire dataset (`X`) to the appropriate key in the dictionary.`model_predictions`.

<!-- grade_id: cell-5030deb704abdd79, points: 0, grade: false, solution: true -->
```python
### GRADED

### YOUR SOLUTION HERE
model_predictions = {f'degree_{i}': None for i in range(1, 11)}

print("Starting Dictionary of Predictions\n", model_predictions)
#for 1, 2, 3, ..., 10

    #create pipeline
    
    #fit pipeline on training data
    
    #make predictions on all data
    
    #assign to model_predictions
    


### BEGIN SOLUTION
X = auto[['horsepower']]
y = auto['mpg']
for i in range(1, 11):
    pipe = Pipeline([('quad_features', PolynomialFeatures(degree = i, include_bias = False)), ('quad_model', LinearRegression())])
    pipe.fit(X_train, y_train)
    preds = pipe.predict(X)
    model_predictions[f'degree_{i}'] = preds
### END SOLUTION

# Answer check
model_predictions['degree_1'][:10]
```

Output:
```
Starting Dictionary of Predictions
 {'degree_1': None, 'degree_2': None, 'degree_3': None, 'degree_4': None, 'degree_5': None, 'degree_6': None, 'degree_7': None, 'degree_8': None, 'degree_9': None, 'degree_10': None}
array([14.90395265,  7.65623939, 10.76240222, 10.76240222, 12.83317743,
        0.82268118, -3.7330243 , -2.69763669, -4.7684119 ,  2.47930135])
```

<!-- grade_id: cell-58b5c109b0fbde13, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
auto_ = pd.read_csv('data/auto.csv')
X_ = auto_[['horsepower']]
y_ = auto_['mpg']
sample_ = auto_.sample(10, random_state = 22)
X_train_ = sample_.loc[:, ['horsepower']]
y_train_ = sample_['mpg']
model_predictions_ = {f'degree_{i}': None for i in range(1, 11)}
for i in range(1, 11):
    pipe_ = Pipeline([('quad_features', PolynomialFeatures(degree = i, include_bias = False)), ('quad_model', LinearRegression())])
    pipe_.fit(X_train_, y_train_)
    preds_ = pipe_.predict(X_)
    model_predictions_[f'degree_{i}'] = preds_
    
#
#
#
#
assert type(model_predictions) == type(model_predictions_)
for i,j in zip(model_predictions.values(), model_predictions_.values()):
    np.testing.assert_array_equal(i, j)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2

#### DataFrame of Predictions

**5 Points**

Use the `model_predictions` dictionary to create a DataFrame of the 10 models predictions.  Assign your solution to `pred_df` below as a DataFrame. 

<!-- grade_id: cell-f410033f079f31d8, points: 0, grade: false, solution: true -->
```python
### GRADED

### YOUR SOLUTION HERE
pred_df = ''
    


### BEGIN SOLUTION
pred_df = pd.DataFrame(model_predictions)
### END SOLUTION

# Answer check
print(type(pred_df))
print(pred_df.head())
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
    degree_1   degree_2   degree_3    degree_4     degree_5     degree_6  \
0  14.903953  14.959892  15.704485   32.550328    97.807527   101.886397   
1   7.656239   9.465786   0.931088 -372.035448 -3456.141665 -4370.275875   
2  10.762402  11.618435   9.428697  -61.767623  -516.945175  -606.298593   
3  10.762402  11.618435   9.428697  -61.767623  -516.945175  -606.298593   
4  12.833177  13.221841  13.121121   13.003201    12.998835    13.007347   

      degree_7     degree_8     degree_9    degree_10  
0   103.934543   103.117944    98.288488    87.834730  
1 -5342.443862 -6208.274949 -6618.861218 -5878.338979  
2  -688.570562  -746.836711  -752.164365  -655.409764  
3  -688.570562  -746.836711  -752.164365  -655.409764  
4    12.999361    12.999488    12.999649    12.999760
```

<!-- grade_id: cell-435b1c581d017378, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
pred_df_ = pd.DataFrame(model_predictions_)
#
#
#
assert type(pred_df_) == type(pred_df)
pd.testing.assert_frame_equal(pred_df, pred_df_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 3

#### DataFrame of Errors

**5 Points**

Now, determine the error for each model and create a DataFrame of these errors.  One way to do this is to use your prediction DataFrame's `.subtract` method to subtract `y` from each feature.  Assign the DataFrame of errors as `error_df` below.  

<!-- grade_id: cell-39bb57c061a5550e, points: 0, grade: false, solution: true -->
```python
### GRADED

### YOUR SOLUTION HERE
error_df = ''
    


### BEGIN SOLUTION
error_df = pred_df.subtract(y, axis = 0)
### END SOLUTION

# Answer check
print(type(error_df))
print(error_df.head())
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
   degree_1  degree_2   degree_3    degree_4     degree_5     degree_6  \
0 -3.096047 -3.040108  -2.295515   14.550328    79.807527    83.886397   
1 -7.343761 -5.534214 -14.068912 -387.035448 -3471.141665 -4385.275875   
2 -7.237598 -6.381565  -8.571303  -79.767623  -534.945175  -624.298593   
3 -5.237598 -4.381565  -6.571303  -77.767623  -532.945175  -622.298593   
4 -4.166823 -3.778159  -3.878879   -3.996799    -4.001165    -3.992653   

      degree_7     degree_8     degree_9    degree_10  
0    85.934543    85.117944    80.288488    69.834730  
1 -5357.443862 -6223.274949 -6633.861218 -5893.338979  
2  -706.570562  -764.836711  -770.164365  -673.409764  
3  -704.570562  -762.836711  -768.164365  -671.409764  
4    -4.000639    -4.000512    -4.000351    -4.000240
```

<!-- grade_id: cell-f8eb586be93a6b46, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
error_df_ = pred_df_.subtract(y, axis = 0)
#
#
#
assert type(error_df_) == type(error_df)
pd.testing.assert_frame_equal(error_df, error_df_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 4

#### Mean and Variance of Model Errors

**5 Points**


Using the DataFrame of errors, examine the mean and variance of each model's error.  What degree model has the highest variance?  Assign your response as an integer to `highest_var_degree` below.

<!-- grade_id: cell-ef9bb67cd8c535da, points: 0, grade: false, solution: true -->
```python
### GRADED

### YOUR SOLUTION HERE
highest_var_degree = ''
    


### BEGIN SOLUTION
highest_var_degree = 10
### END SOLUTION

# Answer check
print(type(highest_var_degree))
print(highest_var_degree)
```

Output:
```
<class 'int'>
10
```

<!-- grade_id: cell-ed6417d67b686a9e, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
highest_var_degree_ = 10
#
#
#
assert highest_var_degree == highest_var_degree_
### END HIDDEN TESTS
```


#### Boxplots of Errors by Degree

Below, uncomment the code to create boxplots for each degree model error.  This should demonstrate an important idea, that as model complexity grows so does the variance in predictions of unseen data. 

```python
# px.box(error_df)
```