## Codio Activity 8.4: Fitting Higher Order Polynomials

**Expected Time: 60 Minutes**

**Total Points: 20**

In this exercise, you will examine the effect of fitting more complex models on the automobile data.  Using the `Pipeline`, you will fit models degree 1 - 10 and evaluate the mean squared error of each model.  Using these results, you are to generate a plot of the model complexity vs. mse.  

#### Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)


```python
import numpy as np
import pandas as pd
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import PolynomialFeatures
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error
import plotly.express as px
```

```python
auto = pd.read_csv('data/auto.csv')
auto.head()
```

Output:
```
mpg  cylinders  displacement  horsepower  weight  acceleration  year  \
0  18.0          8         307.0       130.0    3504          12.0    70   
1  15.0          8         350.0       165.0    3693          11.5    70   
2  18.0          8         318.0       150.0    3436          11.0    70   
3  16.0          8         304.0       150.0    3433          12.0    70   
4  17.0          8         302.0       140.0    3449          10.5    70   

   origin                       name  
0       1  chevrolet chevelle malibu  
1       1          buick skylark 320  
2       1         plymouth satellite  
3       1              amc rebel sst  
4       1                ford torino
```

[Back to top](#Index:) 

### Problem 1

#### Loop of Models

**10 Points** 

Below, construct a loop over the values 1 - 10. For each iteration, fit a regression model with corresponding degree using the `Pipeline`.  Determine the mean squared error and append it to the list `mses` below.  Follow the comments for hints if you have trouble getting started.  Also, be sure to set `include_bias = False` in your transformer.

<!-- grade_id: cell-f0b3f3c52b415e30, points: 0, grade: false, solution: true -->
```python
### GRADED

mses = []
#for 1, 2, 3, ..., 10

    #create pipeline
    
    #fit pipeline
    
    #make predictions
    
    #compute mse
    
    #append mse to mses

### BEGIN SOLUTION
mses = []
X = auto[['horsepower']]
y = auto['mpg']
for i in range(1, 11):
    pipe = Pipeline([('quad_features', PolynomialFeatures(degree = i, include_bias = False)), ('quad_model', LinearRegression())])
    pipe.fit(X, y)
    preds = pipe.predict(X)
    mse = mean_squared_error(y, preds)
    mses.append(mse)
### END SOLUTION

# Answer check
print(len(mses))
print(np.round(mses, 2))
```

Output:
```
10
[23.94 18.98 18.94 18.88 18.43 18.24 18.41 18.57 18.58 18.47]
```

<!-- grade_id: cell-f63f05d9499a5b81, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
mses_ = []
X_ = auto[['horsepower']]
y_ = auto['mpg']
for i in range(1, 11):
    pipe_ = Pipeline([('quad_features', PolynomialFeatures(degree = i, include_bias = False)), ('quad_model', LinearRegression())])
    pipe_.fit(X_, y_)
    preds_ = pipe_.predict(X_)
    mse_ = mean_squared_error(y_, preds_)
    mses_.append(mse_)
#
#
#
assert mses == mses_, 'Make sure you name the steps correctly'
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2

#### Minimum MSE

**10 Points**

Which model complexity had the smallest Mean Squared Error?  Assign your answer as an integer to `best_complexity` below.  

<!-- grade_id: cell-c6e17a14e268d288, points: 0, grade: false, solution: true -->
```python
### GRADED

best_complexity = ''

### BEGIN SOLUTION
best_complexity = mses.index(min(mses)) + 1
### END SOLUTION

# Answer check
print(type(best_complexity))
print(best_complexity)
```

Output:
```
<class 'int'>
6
```

<!-- grade_id: cell-c02091c5b76a08a7, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
mses_ = []
X_ = auto[['horsepower']]
y_ = auto['mpg']
for i in range(1, 11):
    pipe_ = Pipeline([('quad_features', PolynomialFeatures(degree = i, include_bias = False)), ('quad_model', LinearRegression())])
    pipe_.fit(X_, y_)
    preds_ = pipe_.predict(X_)
    mse_ = mean_squared_error(y_, preds_)
    mses_.append(mse_)
best_complexity_ = mses_.index(min(mses_)) + 1
#
#
#
assert best_complexity == best_complexity_, 'Make sure you check the mses list for smallest value.'
### END HIDDEN TESTS
```

### Plotting the Results

Uncomment the code below to generate a plot of the model complexity vs. the mean squared error of your fit pipe.  Do you see the same behavior from the videos -- a decrease in MSE as the complexity increases?

```python
# px.scatter(x = list(range(1, 11)), y = mses,  labels = {'x': 'Degree',
#                                                       'y': "Mean Squared Error"},
#           title = 'Complexity vs. MSE')
```