### Codio Activity 13.4: Decision Boundaries with Two Variables

**Expected Time = 60 minutes** 

**Total Points = 40** 

In the first examples, your work has utilized a straight vertical line as the decision boundary for logistic regression. This is what a decision boundary looks like with only one feature, however with two features the decision boundary becomes a linear function of the two inputs. In this activity, you will focus on generating functions for these boundaries and show strategies for visualizing these boundaries. 

#### Index

  - [Problem 1](#-Problem-1)
  - [Problem 2](#-Problem-2)
  - [Problem 3](#-Problem-3)
  - [Problem 4](#-Problem-4)

```python
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt

from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score, recall_score, precision_score
from sklearn.model_selection import train_test_split
```

### The Data

Again, you will use the penguins data from seaborn.  This time, you will use two features -- `flipper_length_mm` and `bill_length_mm` to build a logistic model and visualize the decision boundary.  The data is loaded and visualized below.  

```python
penguins = sns.load_dataset('penguins').dropna()
penguins = penguins.loc[(penguins['species'] == 'Adelie') | (penguins['species'] == 'Gentoo')]
X = penguins.drop('species', axis = 1)[['flipper_length_mm', 'bill_length_mm']]
y = np.where(penguins.species == 'Adelie', 0, 1)
X_train, X_test, y_train, y_test = train_test_split(X, y, random_state = 42)
```

```python
X.head()
```

```python
sns.scatterplot(data = penguins, x = 'flipper_length_mm', y = 'bill_length_mm', hue = 'species')
```

[Back to top](#-Index)

### Problem 1

#### A Logistic Model

**10 Points**

To get started, instantiate a `LogisticRegression` object as `log_reg` below and fit on the training data.

<!-- grade_id: cell-2a9e2edd23d53626, points: 0, grade: false, solution: true -->
```python
### GRADED

log_reg = ''

### BEGIN SOLUTION
log_reg = LogisticRegression().fit(X_train, y_train)
### END SOLUTION

# Answer check
print(log_reg)
```

<!-- grade_id: cell-5edb70583e600596, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
log_reg_ = LogisticRegression().fit(X_train, y_train)
#
#
#
assert type(log_reg) == type(log_reg_)
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### Examining the coefficients

**10 Points**

Now, our sigma function looks like:

$$\sigma(x) = \frac{1}{1 + e^{-z}}$$

where

$$z = \beta_0 + \beta_1 * x_0 + \beta_2 * x_1.$$

Below, assign the coefficients and intercept to `beta_0`, `beta_1` and `beta_2` respectively as floats.  Note that $x_0$ is flipper length.  

<!-- grade_id: cell-de41271c6395179d, points: 0, grade: false, solution: true -->
```python
### GRADED

beta_0 = ''
beta_1 = ''
beta_2 = ''

### BEGIN SOLUTION
beta_0 = float(log_reg.intercept_)
beta_1 = float(log_reg.coef_[0][0])
beta_2 = float(log_reg.coef_[0][1])
### END SOLUTION

# Answer check
print(f'z = {beta_0: .2f} + {beta_1: .2f}x0 + {beta_2: .2f}x1')
```

<!-- grade_id: cell-ba95f7318397f251, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
log_reg_ = LogisticRegression().fit(X_train, y_train)
beta_0_ = float(log_reg_.intercept_)
beta_1_ = float(log_reg_.coef_[0][0])
beta_2_ = float(log_reg_.coef_[0][1])
#
#
#
assert beta_0 == beta_0_
assert beta_1 == beta_1_
assert beta_2 == beta_2_
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Visualizing the decision boundary

**10 Points**

<center>
    <img src = 'images/dboundary.png' />
</center>


There is both a brute force and more formal approach for visualizing the decision boundary.  What we won't do, is create a grid of points, use the model to make predictions for them, and plot colored regions according to their predicted value. 

Instead, with two variables we can directly solve for the linear function in terms of $x_0$ and $x_1$, and your $\beta$'s. Upon doing so we find a linear function defined as: 

$$y = -\frac{\beta_1}{\beta_2} * x_0 - \frac{\beta_0}{\beta_2}$$

Complete the function `decision_boundary` below that takes in $x_0$ and returns the appropriate value for $y$ based on this formula.  Uncomment the plot to verify your results using the defined `x = np.linspace`.

<!-- grade_id: cell-7e6d6568117a8e6b, points: 0, grade: false, solution: true -->
```python
### GRADED

def decision_boundary(x0, beta_0, beta_1, beta_2):
    '''
    Function returns values for linear decision boundaries
    in binary classification setting accoring to the formula
    y = -beta_1/beta_2 * x0 - beta_0/beta_2
    
    Arguments
    ---------
    x0: np.array
        domain for evaluation of function
        
    beta_0: float
         intercept from fit logistic model
    beta_1: float
        first coefficient from logistic model
    beta_2: float
        second coefficient from logistic model
    
    Returns
    -------
    np.array
        values of y
    '''
    return None
    
### BEGIN SOLUTION
def decision_boundary(x0, beta_0, beta_1, beta_2):
    return -(beta_1/beta_2)*x0 - beta_0/beta_2
### END SOLUTION
x = np.linspace(165, 240, 100)
print(decision_boundary(x, beta_0, beta_1, beta_2)[0], decision_boundary(x, beta_0, beta_1, beta_2)[-1])
```

```python
## code for figure
# x = np.linspace(165, 240, 100)
# sns.scatterplot(data = penguins, x = 'flipper_length_mm', y = 'bill_length_mm', hue = 'species')
# plt.plot(x, decision_boundary(x, beta_0, beta_1, beta_2), '--', color = 'black')
# plt.ylim(25, 65)
# plt.fill_between(x, decision_boundary(x, beta_0, beta_1, beta_2), alpha = 0.3, color = 'lightblue')
# plt.fill_between(x, decision_boundary(x, beta_0, beta_1, beta_2), np.repeat(70, 100), alpha = 0.3)
```

<!-- grade_id: cell-6713f398a08b64d4, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def decision_boundary_(x0, beta_0, beta_1, beta_2):
    return -(beta_1/beta_2)*x0 - beta_0/beta_2
#
#
#
np.testing.assert_array_equal(decision_boundary(x, beta_0, beta_1, beta_2),decision_boundary_(x, beta_0, beta_1, beta_2))
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 4

#### Comparing regressors

**10 Points**

Now, fit a second regressor using the argument `C = 0.001`. Compare the decision boundary by using what you've seen earlier. How did the decision boundary change based on this?  The slope of the new decision boundary should either be more or less than the default settings.  Assign your answer as a string -- `greater than` or `less than` to `ans4` below. Remember that the slope is negative. 



<!-- grade_id: cell-66dfc51fc45c7a0f, points: 0, grade: false, solution: true -->
```python
### GRADED

ans4 = ''

### BEGIN SOLUTION
ans4 = 'less than'
### END SOLUTION

# Answer check
print(ans4)
```

<!-- grade_id: cell-d0c153c049b09fc1, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans4_ = 'less than'
#
#
#
assert ans4 == ans4_
### END HIDDEN TESTS
```